This notebook will be collected automatically at **6pm on Monday** from `/home/data_scientist/assignments/Week9` directory on the course JupyterHub server. If you work on this assignment on the course Jupyterhub server, just make sure that you save your work and instructors will pull your notebooks automatically after the deadline. If you work on this assignment locally, the only way to submit assignments is via Jupyterhub, and you have to place the notebook file in the correct directory with the correct file name before the deadline.

1. Make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).
2. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.
3. Do not change the file path or the file name of this notebook.
4. Make sure that you save your work (in the menubar, select `File` → `Save and CheckPoint`)

# Problem 9.3. NLP: Semantic Analysis.

In this problem, we explore semantic analysis.

```python
import string
import numpy as np
from nltk.tokenize import WordPunctTokenizer
from sklearn.datasets import load_files
import gensim

from nose.tools import (
    assert_equal,
    assert_is_instance,
    assert_almost_equal,
    assert_true
    )
```

## Wordnet

We use the Wordnet synonym rings.

```python
from nltk.corpus import wordnet as wn
```

- Find how many entries a word has in the wordnet synset.

<!-- grade_id: find_number_of_entries_in_synonym_ring_answer, points: 0, grade: false, solution: true -->
```python
def find_number_of_entries_in_synonym_ring(word):
    '''
    Finds the number of entries in the wordnet synset.
    
    Parameters
    ----------
    word: A string.
    
    Returns
    -------
    An int.
    '''
    
    # YOUR CODE HERE
    
    return result
```

```python
the_word = 'ship'
n_entries = find_number_of_entries_in_synonym_ring(the_word)
print('{0} total entries in synonym ring for {1}. '.format(n_entries, the_word))
```

```python
the_word = 'throw'
n_entries = find_number_of_entries_in_synonym_ring(the_word)
print('{0} total entries in synonym ring for {1}. '.format(n_entries, the_word))
```

<!-- grade_id: find_number_of_entries_in_synonym_ring_test, points: 3, grade: true, solution: false -->
```python
assert_is_instance(find_number_of_entries_in_synonym_ring('ship'), int)
assert_equal(find_number_of_entries_in_synonym_ring('ship'), 6)
assert_equal(find_number_of_entries_in_synonym_ring('throw'), 20)
```

## Word Similarities

- Compute the path similarity for _cat_, _dog_, _boy_, and _girl_.

<!-- grade_id: get_path_similarity_between_boy_and_girl_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_boy_and_girl():
    '''
    Computes the path similarity between "boy" and "girl".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: get_path_similarity_between_boy_and_girl_test, points: 2, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_boy_and_girl(), float)
assert_almost_equal(get_path_similarity_between_boy_and_girl(), 0.16666666666666666)
```

<!-- grade_id: get_path_similarity_between_boy_and_cat_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_boy_and_cat():
    '''
    Computes the path similarity between "boy" and "cat".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: get_path_similarity_between_boy_and_cat_test, points: 2, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_boy_and_cat(), float)
assert_almost_equal(get_path_similarity_between_boy_and_cat(), 0.08333333333333333)
```

<!-- grade_id: get_path_similarity_between_boy_and_dog_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_boy_and_dog():
    '''
    Computes the path similarity between "boy" and "dog".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: get_path_similarity_between_boy_and_dog_test, points: 2, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_boy_and_dog(), float)
assert_almost_equal(get_path_similarity_between_boy_and_dog(), 0.14285714285714285)
```

<!-- grade_id: get_path_similarity_between_girl_and_girl_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_girl_and_girl():
    '''
    Computes the path similarity between "girl" and "girl".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: get_path_similarity_between_girl_and_girl_test, points: 2, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_girl_and_girl(), float)
assert_almost_equal(get_path_similarity_between_girl_and_girl(), 1.0)
```

```python
# Now we print similarity measures.
fmt_str = '{1} to {2}: {0:4.3f}'

print('Path Similarity:')
print(40*'-')
print(fmt_str.format(get_path_similarity_between_boy_and_girl(), 'boy', 'girl'))
print(fmt_str.format(get_path_similarity_between_boy_and_cat(), 'boy', 'cat'))
print(fmt_str.format(get_path_similarity_between_boy_and_dog(), 'boy', 'dog'))
print(fmt_str.format(get_path_similarity_between_girl_and_girl(), 'girl', 'girl'))
```

## Word2Vec

In the second half of this problem, we use the NLTK Brown corpus to build a word2vec model.

```python
from nltk.corpus import brown
sentences = brown.sents()
```

### Word2Vec model

- Build a `Word2Vec` model from movie reviews `new_mvr`.
- Set the maximum distance between the current and predicted word within a sentence to 10.
- Ignore all words with total frequency lower than 2.

<!-- grade_id: get_model_answer, points: 0, grade: false, solution: true -->
```python
def get_model(sentences):
    '''
    Builds a Word2Vec model from "corpus"
    
    Parameters
    ----------
    corpus: An NTLK corpus.
    
    Returns
    -------
    A Word2Vec instance.
    '''
    
    # YOUR CODE HERE

    return model
```

The following code cell takes a while to complete.

```python
model = get_model(sentences)
```

<!-- grade_id: get_model_test, points: 2, grade: true, solution: false -->
```python
assert_is_instance(model, gensim.models.Word2Vec)
assert_equal(model.window, 10)
assert_equal(model.min_count, 2)
```

## Cosine Similarity

- Compute Cosine Similarities.

<!-- grade_id: get_cosine_similarity_answer, points: 0, grade: false, solution: true -->
```python
def get_cosine_similarity(model, word1, word2):
    '''
    Computes cosine similarity between "word1" and "word2" using a Word2Vec model.
    
    Parameters
    ----------
    model: A gensim.Word2Vec model.
    
    Returns
    -------
    A float
    '''
    
    # YOUR CODE HERE
    
    return similarity
```

```python
# Now we print similarity measures.
fmt_str = '{1} to {2}: {0:4.3f}'

print('Cosine Similarity:')
print(40*'-')
print(fmt_str.format(get_cosine_similarity(model, 'boy', 'girl'), 'boy', 'girl'))
print(fmt_str.format(get_cosine_similarity(model, 'boy', 'cat'), 'boy', 'cat'))
print(fmt_str.format(get_cosine_similarity(model, 'boy', 'dog'), 'boy', 'dog'))
print(fmt_str.format(get_cosine_similarity(model, 'girl', 'girl'), 'girl', 'girl'))
```

<!-- grade_id: get_cosine_similarity_test, points: 3, grade: true, solution: false -->
```python
assert_is_instance(get_cosine_similarity(model, 'boy', 'girl'), float)
assert_almost_equal(get_cosine_similarity(model, 'boy', 'girl'), model.similarity('boy', 'girl'))

assert_almost_equal(get_cosine_similarity(model, 'boy', 'cat'), model.similarity('boy', 'cat'))
assert_almost_equal(get_cosine_similarity(model, 'boy', 'dog'), model.similarity('boy', 'dog'))
assert_almost_equal(get_cosine_similarity(model, 'girl', 'girl'), model.similarity('girl', 'girl'))
```

## Most similar words

- Find the top 3 most similar words, where "girl" and "cat" contribute positively towards the similarity, and "boy" and "dog" contribute negatively.

<!-- grade_id: find_most_similar_words_answer, points: 0, grade: false, solution: true -->
```python
def find_most_similar_words(model):
    '''
    Find the top 3 most similar words,
    where "girl" and "cat" contribute positively towards the similarity,
    and "boy" and "dog" contribute negatively.
    
    Parameters
    ----------
    model: A gensim.Word2Vec model.
    
    Returns
    -------
    A list of tuples (word, similarty).
    word: A string.
    similarity: A float.
    '''
    
    # YOUR CODE HERE
    
    return result

```

```python
print('{0:14s}: {1}'.format('Word', 'Cosine Similarity'))
print(40*'-')
for val in find_most_similar_words(model):
    print('{0:14s}: {1:6.3f}'.format(val[0], val[1]))
```

<!-- grade_id: find_most_similar_words_test, points: 4, grade: true, solution: false -->
```python
assert_is_instance(find_most_similar_words(model), list)
assert_true(all(isinstance(t[0], str) for t in find_most_similar_words(model)))
assert_true(all(isinstance(t[1], float) for t in find_most_similar_words(model)))
assert_equal(len(find_most_similar_words(model)), 3)
words = [t[0] for t in model.most_similar(positive=['girl', 'cat'], negative=['boy', 'dog'], topn=3)]
similarities = [t[1] for t in model.most_similar(positive=['girl', 'cat'], negative=['boy', 'dog'], topn=3)]
assert_equal([t[0] for t in find_most_similar_words(model)], words)
assert_almost_equal(find_most_similar_words(model)[0][1], similarities[0])
assert_almost_equal(find_most_similar_words(model)[1][1], similarities[1])
assert_almost_equal(find_most_similar_words(model)[2][1], similarities[2])
```