This notebook will be collected automatically at **6pm on Monday** from `/home/data_scientist/assignments/Week3` directory on the course JupyterHub server. If you work on this assignment on the course Jupyterhub server, just make sure that you save your work and instructors will pull your notebooks automatically after the deadline. If you work on this assignment locally, the only way to submit assignments is via Jupyterhub, and you have to place the notebook file in the correct directory with the correct file name before the deadline.

1. Make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).
2. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.
3. Do not change the file path or the file name of this notebook.
4. Make sure that you save your work (in the menubar, select `File` → `Save and CheckPoint`)

# Problem 3.3. Supervised Learning: Naive Bayes

In this problem, we will implement Naive Bayes from scratch using **Numpy**.

```python
import numpy as np
from numpy.testing import assert_array_equal, assert_array_almost_equal
```

## Separate Data By Class

Write a function that takes two Numpy arrays, separates data by class value, and returns the result as a dictionary. The keys of the dictionary are class values, and the dictionary values are the rows in `X` that correspond to each class value.

<!-- grade_id: separate_by_class_answer, points: 0, grade: false, solution: true -->
```python
def separate_by_class(X, y):
    '''
    Separate the training set ("X") by class value ("y")
    so that we can calculate statistics for each class.
    
    Parameters
    ----------
    X: A 2d numpy array
    y: A 1d numpy array
    Returns
    -------
    A dictionary of 2d numpy arrays
    '''
    ##################
    # YOUR CODE HERE
    ##################

    return separated
```

<!-- grade_id: separate_by_class_test, points: 3, grade: true, solution: false -->
```python
X_t = np.array( [[2, 21], [1, 20], [3, 22]] )
y_t = np.array( [1, 0, 1] )
separated_t = separate_by_class(X_t, y_t)
assert_array_equal(separated_t[0], np.array( [ [1, 20] ] ))
assert_array_equal(separated_t[1], np.array( [ [2, 21], [3, 22] ] ))
```

## Calculate Mean

We calculate the mean and use it as the middle of our gaussian distribution when calculating probabilities. If the input array is a 2d array (i.e., a matrix), you should take the mean of each **column**.

<!-- grade_id: calculate_mean_answer, points: 0, grade: false, solution: true -->
```python
def calculate_mean(array):
    '''
    Calculates the mean of each column, i.e. each attribute.
    
    Parameters
    ----------
    A 1d or 2d numpy array
    
    Returns
    -------
    A 1d or 2d numpy array
    '''
    ##################
    # YOUR CODE HERE
    ##################
    return mean
```

<!-- grade_id: calculate_mean_test, points: 3, grade: true, solution: false -->
```python
array_t = np.array( [ [1, 4, 7], [2, 5, 6], [3, 6, 8] ] )
mean_t = calculate_mean(array_t)
assert_array_equal(mean_t, np.array( [2., 5., 7.] ))
```

## Calculate Standard Deviation

Write a function that calculates the standard deviation of each **column**. The input array may be a 2d array.

<!-- grade_id: calculate_stdev_answer, points: 0, grade: false, solution: true -->
```python
def calculate_stdev(array):
    '''
    Calculates the standard deviation of each column, i.e. each attribute.

    Parameters
    ----------
    A 1d or 2d numpy array
    
    Returns
    -------
    A 1d or 2d numpy array
    '''
    ##################
    # YOUR CODE HERE
    ##################

    return stdev
```

<!-- grade_id: calculate_stdev_test, points: 3, grade: true, solution: false -->
```python
array_t = np.array( [ [1, 20, 14], [2, 21, 15], [3, 22, 16] ] )
stdev_t = calculate_stdev(array_t)
assert_array_equal(stdev_t, np.array( [1., 1., 1.] ))
```

## Summarize data set

For a given list of instances (for a class value), we calculate the mean and the standard deviation for each attribute.

<!-- grade_id: summarize_answer, points: 0, grade: false, solution: true -->
```python
def summarize(X):
    '''
    For a given list of instances (for a class value),
    calculates the mean and the standard deviation for each attribute.
    
    Parameters
    ----------
    A 2d numpy array
    
    Returns
    -------
    A 2d numpy array
    '''
    ##################
    # YOUR CODE HERE
    ##################
    return summary
```

<!-- grade_id: summarize_test, points: 3, grade: true, solution: false -->
```python
X_t = np.array( [ [1, 20], [2, 21], [3, 22] ] )
summary_t = summarize(X_t)
assert_array_equal(summary_t, np.array( [ (2.0, 1.0), (21.0, 1.0) ] ))
```

## Summarize Attributes By Class

We calculate the summaries for each attribute.

<!-- grade_id: summarize_by_class_answer, points: 0, grade: false, solution: true -->
```python
def summarize_by_class(X, y):
    '''
    Separates a training set into instances grouped by class.
    It then calculates the summaries for each attribute.
    
    Parameters
    ----------
    X: A 2d numpy array. Represents training attributes.
    y: A 1d numpy array. Represents class labels.
    Returns
    -------
    A dictionary of 2d numpy arrays
    '''
    ##################
    # YOUR CODE HERE
    ##################

    return summaries
```

<!-- grade_id: summarize_by_class_test, points: 3, grade: true, solution: false -->
```python
X_t = np.array( [ [1, 20], [2, 21], [3, 22], [4, 22] ] )
y_t = np.array( [1, 0, 1, 0] )
summary_t = summarize_by_class(X_t, y_t)
assert_array_almost_equal(summary_t[0], np.array( [ (3., 1.41421356), (21.5, 0.70710678) ] ))
assert_array_almost_equal(summary_t[1], np.array( [ (2., 1.41421356), (21.0, 1.41421356) ] ))
```

## Calculate Log of Gaussian Probability Density Function

Calculate the **log** of a Gaussian Probability Density Function. The conditional probabilities for each class given an attribute value are small. When they are multiplied together they result in very small values, which can lead to floating point underflow (numbers too small to represent in Python). A common fix for this is to combine the log of the probabilities together.

<!-- grade_id: calculate_log_probability_answer, points: 0, grade: false, solution: true -->
```python
def calculate_log_probability(x, mean, stdev):
    '''
    Calculates log of Gaussian function to estimate
    the log probability of a given attribute value.
    
    Parameters
    ----------
    x: A float or 1d numpy array
    mean: A float or 1d numpy array
    stdev: A float or 1d numpy array
    
    Returns
    -------
    A float or 1d numpy array
    '''

    ##################
    # YOUR CODE HERE
    ##################

    return log_probability
```

<!-- grade_id: calculate_log_probability_test, points: 3, grade: true, solution: false -->
```python
array_t = calculate_log_probability(np.array( [71.5] ), np.array( [73] ), np.array( [6.2] ))
assert_array_almost_equal(array_t, np.array( [ np.log(0.0624896575937) ] ))
```

## Calculate Class Probabilities

Remember that you calculated **log** of probabilities. Therefore, instead of combine probabilities together by multiplying them, you should **add** the log of probabilities.

<!-- grade_id: calculate_class_log_probabilities_answer, points: 0, grade: false, solution: true -->
```python
def calculate_class_log_probabilities(summaries, input_array):
    '''
    Combines the probabilities of all of the attribute values for a data instance
    and comes up with a probability of the entire data instance belonging to the class.

    Parameters
    ----------
    summaries: A dictionary of 2d numpy arrays
    input_array: A 1d numpy array
    
    Returns
    -------
    A dictionary of log probabilities
    '''

    ##################
    # YOUR CODE HERE
    ##################

    return log_probabilities
```

<!-- grade_id: calculate_class_log_probabilities_test, points: 6, grade: true, solution: false -->
```python
summaries_t = {0: np.array( [ (1, 0.5) ]), 1: np.array( [ (20, 5.0) ] )}
input_t = np.array( [1.1] )
log_probabilities = calculate_class_log_probabilities(summaries_t, input_t)
assert_array_almost_equal(log_probabilities[0], np.array( [np.log(0.7820853879509118)] ))
assert_array_almost_equal(log_probabilities[1], np.array( [np.log(6.298736258150442e-05)] ))
```

## Make Predictions

Calculate the probability of a data instance belonging to each class value, we can look for the largest probability and return the associated class.

<!-- grade_id: predict_answer, points: 0, grade: false, solution: true -->
```python
def predict(summaries, input_array):
    '''
    Calculates the probability of each data instance belonging to each class value,
    looks for the largest probability, and return the associated class.
    
    Parameters
    ----------
    summaries: A dictionary of numpy arrays
    input_array: A 1d numpy array
    
    Returns
    -------
    A 1d numpy array
    '''

    ##################
    # YOUR CODE HERE
    ##################

    return best_label
```

<!-- grade_id: predict_test, points: 6, grade: true, solution: false -->
```python
summaries_t = {0: np.array( [ (1, 0.5) ] ), 1: np.array( [ (20, 5.0) ] )}
input_t1 = np.array( [1.1] )
result_t1 = predict(summaries_t, input_t1)
assert_array_equal(result_t1, np.array( [0.] ))

test_set_t2 = np.array([1.1, 19.1])
result_t2 = predict(summaries_t, test_set_t2)
assert_array_equal(result_t2, np.array( [0., 1.] ))
```

If you used Numpy correctly, it shouldn't be necessary to iterate over each test instance and make predictions. In other words, you don't have to use the `for` loop of `getPredictions()`.