This notebook will be collected automatically at **6pm on Monday** from `/home/data_scientist/assignments/Week13` directory on the course JupyterHub server. If you work on this assignment on the course Jupyterhub server, just make sure that you save your work and instructors will pull your notebooks automatically after the deadline. If you work on this assignment locally, the only way to submit assignments is via Jupyterhub, and you have to place the notebook file in the correct directory with the correct file name before the deadline.

1. Make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).
2. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.
3. Do not change the file path or the file name of this notebook.
4. Make sure that you save your work (in the menubar, select `File` → `Save and CheckPoint`)

## Problem 13.1. MongoDB

In this problem, we work with MongoDB from a Python program by using the pymongo database driver.

```python
import os
import datetime
import json
import bson
import pymongo as pm

from nose.tools import assert_equal, assert_true, assert_is_instance
```

We use the historical weather data from [Weather Underground](http://www.wunderground.com/) to create a database. For more information on this data set, see [Week 4 Problem 1](https://github.com/UI-DataScience/info490-sp16/blob/master/Week4/assignments/w4p1.ipynb).

```python
fpath = '/home/data_scientist/data/weather'
fname = 'weather_kord_2001_0101.json'

with open(os.path.join(fpath, fname)) as f:
    weather_json = json.load(f)

assert_is_instance(weather_json, dict)
assert_equal(set(weather_json.keys()), set(['current_observation', 'response', 'history']))
assert_true('observations' in weather_json['history'])
```

```python
observations = weather_json['history']['observations']
print('There are {} dictionaries in the list.'.format(len(observations)))
print('The first element is\n{}'.format(observations[0]))

assert_is_instance(observations, list)
assert_true(all(isinstance(o, dict) for o in observations))
```

We connect to the course MongoDB cloud computing system, hosted by Microsoft Azure.

```python
client = pm.MongoClient("mongodb://104.214.79.74:27017")
```

Since we are using a shared resource without authentication, we use your netid to create a database for each student.

```python
# Filename containing user's netid
fname = '/home/data_scientist/users.txt'
with open(fname, 'r') as fin:
    netid = fin.readline().rstrip()

# We will delete our working directory if it exists before recreating.
dbname = 'assignment-{0}'.format(netid)

if dbname in client.database_names():
    client.drop_database(dbname)

print('Existing databases:', client.database_names())

assert_true(dbname not in client.database_names())
```

## Inserting Data

- Create a new collection using the name `collection_name` and add new documents `data` to our MongoDB collection
- Return a list of object IDs as a validation of the insertion process.

<!-- grade_id: insert_answer, points: 0, grade: false, solution: true -->
```python
def insert_data(db, collection_name, data):
    '''
    Creates a new collection using the name "collection_name" 
    and adds new documents `data` to our MongoDB collection.
    
    Parameters
    ----------
    data: A list of dictionaries.
    db: A pymongo.database.Database instance.
    collection_name: Name of new MongoDB collection.
    
    Returns
    -------
    A list of bson.ObjectId
    '''
    
    # YOUR CODE HERE
    
    return inserted_ids
```

```python
inserted_ids = insert_data(client[dbname], '0101', observations)

print("New weather ID: ", inserted_ids)
print('Existing databases:', client.database_names())
print('Existing collections:', client[dbname].collection_names())
```

<!-- grade_id: insert_test, points: 6, grade: true, solution: false -->
```python
assert_is_instance(inserted_ids, list)
assert_true(all(isinstance(i, bson.objectid.ObjectId) for i in inserted_ids))

assert_true(dbname in client.database_names())
assert_true('0101' in client[dbname].collection_names())
assert_equal(client[dbname]['0101'].count(), len(observations))
```

## Retrieving Data

- Find all documents that have a given weather `condition` (e.g., `conds == "Clear"` or `conds == "Partly Cloudy"`)
- Return the `_id` values of all documents that match the search query.

<!-- grade_id: retrieve_answer, points: 0, grade: false, solution: true -->
```python
def retrieve_data(collection, condition):
    '''
    Finds all documents that have a given weather `condition`
    and return the `_id` values of all documents that match the search query.
    
    Parameters
    ----------
    collection: A pymongo.Collection instance.
    condition: A string, e.g., "Clear", "Partly Cloudy", "Overcast".
    
    Returns
    -------
    A list of bson.ObjectId
    '''
    
    # YOUR CODE HERE
    
    return result
```

```python
clear_ids = retrieve_data(client[dbname]['0101'], 'Clear')
print(clear_ids)
```

<!-- grade_id: retrieve_test, points: 6, grade: true, solution: false -->
```python
assert_is_instance(clear_ids, list)
assert_true(all(isinstance(i, bson.objectid.ObjectId) for i in clear_ids))

conds = {obs['conds'] for obs in observations}
for cond in conds:
    r = retrieve_data(client[dbname]['0101'], cond)
    n = [obs['_id'] for obs in observations if obs['conds'] == cond]
    assert_equal(len(r), len(n))
    assert_equal(set(r), set(n))
```

## Modifying Data

- Find all documents whose `conds` value is `"Clear"` and change the `conds` attribute to `Cloudy`.
- Return the number of documents modified as a validation of the process.

<!-- grade_id: modify_answer, points: 0, grade: false, solution: true -->
```python
def modify_data(collection):
    '''
    Finds all documents whose "conds" value is "Clear"
    and change the "conds" attribute to "Cloudy".

    Parameters
    ----------
    collection: A pymongo.Collection instance.
    
    Returns
    -------
    An int. The number of documents modified.
    '''
    
    # YOUR CODE HERE
    
    return count
```

```python
n_modified = modify_data(client[dbname]['0101'])
print('{0} records modified.'.format(n_modified))
```

<!-- grade_id: modify_test, points: 6, grade: true, solution: false -->
```python
assert_equal(
    n_modified,
    len([obs['_id'] for obs in observations if obs['conds'] == 'Clear'])
    )

conds = [obs['conds'] for obs in observations]
for cond in conds:
    if cond != 'Clear' and cond != 'Cloudy':
        r = retrieve_data(client[dbname]['0101'], cond)
        n = [obs['_id'] for obs in observations if obs['conds'] == cond]
        assert_equal(len(r), len(n))
        assert_equal(set(r), set(n))
```

## Advanced Querying

- Find all documents with `visi` equal to `"10.0"` and sort the documents by `conds`.
- Return a list of `conds` as a validation of the process.

<!-- grade_id: query_answer, points: 0, grade: false, solution: true -->
```python
def query(collection):
    '''
    Finds all documents with "visi" equal to `"10.0"
    and sort the documents by "conds".
    
    Parameters
    ----------
    collection: A pymongo.Collection instance.

    Returns
    -------
    A list of strings.
    '''
    
    # YOUR CODE HERE

    return result
```

```python
query_conds = query(client[dbname]['0101'])
print(query_conds)
```

<!-- grade_id: query_test, points: 6, grade: true, solution: false -->
```python
modified_conds = [obs['conds'] for obs in observations if obs['visi'] == '10.0']
modified_conds = ['Cloudy' if cond == 'Clear' else cond for cond in modified_conds]
modified_conds = sorted(modified_conds)
assert_equal(query_conds, modified_conds)
```

## Deleting Data

- Delete all documents whose `conds` attribute is equal to `"Cloudy"` from our collection.
- Return the number of documents deleted as a validation of the process.

<!-- grade_id: delete_answer, points: 0, grade: false, solution: true -->
```python
def delete_data(collection):
    '''
    Deletes all documents whose "conds" == "Cloudy".
    
    Paramters
    ---------
    collection: A pymongo.Collection instance.

    Returns
    -------
    An int. The number of documents deleted.
    '''
    
    # YOUR CODE HERE
    
    return count
```

```python
n_deleted = delete_data(client[dbname]['0101'])
print('{0} records deleted.'.format(n_deleted))
```

<!-- grade_id: delete_test, points: 6, grade: true, solution: false -->
```python
deleted_obs = [obs for obs in modified_conds if obs == 'Cloudy']
assert_equal(n_deleted, len(deleted_obs))

for cond in set(conds):
    if cond != 'Clear' and cond != 'Cloudy':
        r = retrieve_data(client[dbname]['0101'], cond)
        n = [obs['_id'] for obs in observations if obs['conds'] == cond]
        assert_equal(len(r), len(n))
        assert_equal(set(r), set(n))
```

## Cleanup

When you are done or if you want to start over with a clean database, run the following code cell.

```python
if dbname in client.database_names():
    client.drop_database(dbname)
    
assert_true(dbname not in client.database_names())
```