This notebook will be collected automatically at **6pm on Monday** from `/home/data_scientist/assignments/Week8` directory on the course JupyterHub server. If you work on this assignment on the course Jupyterhub server, just make sure that you save your work and instructors will pull your notebooks automatically after the deadline. If you work on this assignment locally, the only way to submit assignments is via Jupyterhub, and you have to place the notebook file in the correct directory with the correct file name before the deadline.

1. Make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).
2. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.
3. Do not change the file path or the file name of this notebook.
4. Make sure that you save your work (in the menubar, select `File` → `Save and CheckPoint`)

# Problem 8.1. Social Media: Email

In this problem, we first explore parsing a simple email, and develop a text classification pipeline from a public email corpus.

```python
import os
import email as em
from email import policy
import numpy as np
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.svm import LinearSVC
from sklearn.pipeline import Pipeline
from sklearn.metrics import accuracy_score
from sklearn.utils import check_random_state

from nose.tools import assert_equal, assert_is_instance, assert_true
from numpy.testing import assert_array_equal, assert_almost_equal
```

## Email header

- Write a function that extracts the header information (`to`, `from`, and `subject` fields) from a text file.

<!-- grade_id: read_header_answer, points: 0, grade: false, solution: true -->
```python
def read_header(filename):
    '''
    Extracts email header from a text file.
    
    Parameters
    ----------
    filename: A string.
    
    Returns
    -------
    A 3-tuple of (msg_to, msg_from, subject).
    
    msg_to: A string. The "To" field of the email header.
    msg_from: A string. The "From" field.
    subject: A string. The "subject".
    '''
    
    # YOUR CODE HERE
    
    return msg_to, msg_from, subject
```

```python
email1 = '/home/data_scientist/data/email/ham/00001.1a31cc283af0060967a233d26548a6ce'
to1, from1, subject1 = read_header(email1)
print('To:', to1)
print('From:', from1)
print('Subject:', subject1)
```

<!-- grade_id: reader_header_test_1, points: 3, grade: true, solution: false -->
```python
assert_equal(to1, 'Chris Garrigues <cwg-dated-1030314468.7c7c85@DeepEddy.Com>')
assert_equal(from1, 'Robert Elz <kre@munnari.OZ.AU>')
assert_equal(subject1, 'Re: New Sequences Window')
```

```python
email2 = '/home/data_scientist/data/email/spam/00001.317e78fa8ee2f54cd4890fdc09ba8176'
to2, from2, subject2 = read_header(email2)
print('To:', to2)
print('From:', from2)
print('Subject:', subject2)
```

<!-- grade_id: reader_header_test_2, points: 2, grade: true, solution: false -->
```python
assert_equal(to2, 'ilug@linux.ie')
assert_equal(from2, 'Start Now <startnow2002@hotmail.com>')
assert_equal(subject2, '[ILUG] STOP THE MLM INSANITY')
```

## Payload

- Write a function named `get_payload()` that iterates through each file in a directory specified by the `path` parameter and returns a list of emails. The function only reads in files in a specified range. For example, if `start` is 500 and `end` is 999, the function reads all emails from the 500th email (in the order as determined by `os.walk()`) to the 999th email, inclusive.

<!-- grade_id: get_payload_answer, points: 0, grade: false, solution: true -->
```python
def get_payload(path, start, end):
    '''
    Iterates through each file in a directory and reads emails.
    
    Paremters
    ---------
    path: A string.
    start: An int.
    end: An int.
    
    Returns
    -------
    A list of strings.
    '''
    
    # YOUR CODE HERE
    
    return result
```

```python
ham = get_payload('/home/data_scientist/data/email/ham', start=500, end=999)
print(ham[0])
```

<!-- grade_id: get_payload_test_1, points: 5, grade: true, solution: false -->
```python
assert_is_instance(ham, list)
assert_true(all(isinstance(h, str) for h in ham))
assert_equal(len(ham), 500)
assert_true(ham[0].startswith('Ad Aware sometimes helps too.'))
assert_true(ham[-1].startswith('On 10 Aug 2002, Gary Lawrence Murphy wrote:'))
```

```python
spam = get_payload('/home/data_scientist/data/email/spam', start=500, end=999)
print(spam[0])
```

<!-- grade_id: get_payload_test_2, points: 5, grade: true, solution: false -->
```python
assert_is_instance(spam, list)
assert_true(all(isinstance(s, str) for s in spam))
assert_equal(len(spam), 500)
assert_true(spam[0].startswith('Dear Sirs,\nWe know your esteemed company in beach towels'))
assert_true(spam[-1].startswith('Hi,\n\nI\'m a college dropout.'))
```

## Convert list to Numpy array

- Convert `ham` and `spam` to Numpy arrays and return them as `pos_emails` and `neg_emails`, respectively. Return Numpy arrays for the labels, where each `ham` is labeled 1 and each `spam` 0.

<!-- grade_id: list_to_array_answer, points: 0, grade: false, solution: true -->
```python
def list_to_array(ham, spam):
    '''
    Converts a list of emails to Numpy array and labels them.
    
    Parameters
    ----------
    ham: A list of strings.
    spam: A lsit of strings.
    
    Returns
    -------
    A 4-tuple of (pos_emails, neg_emails, pos_labels, neg_labels)
    All Numpy arrays.
    '''
    
    # YOUR CODE HERE
    
    return pos_emails, neg_emails, pos_labels, neg_labels
```

```python
pos_emails, neg_emails, pos_labels, neg_labels = list_to_array(ham, spam)
```

<!-- grade_id: list_to_array_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(pos_emails, np.ndarray)
assert_is_instance(neg_emails, np.ndarray)
assert_is_instance(pos_labels, np.ndarray)
assert_is_instance(neg_labels, np.ndarray)
assert_array_equal(pos_emails, ham)
assert_array_equal(neg_emails, spam)
assert_array_equal(pos_labels, [1] * len(ham))
assert_array_equal(neg_labels, [0] * len(spam))
```

## Create training and test sets

- Create four new arrays from our email data and labels.

<!-- grade_id: train_test_split_answer, points: 0, grade: false, solution: true -->
```python
def train_test_split(pos_emails, neg_emails, pos_labels, neg_labels, split_value):
    '''
    Creates four arrays from emails and labels by using "split_value"
    
    Parameters
    ----------
    pos_emails: A numpy array of strings.
    neg_emails: A numpy array of strings.
    pos_labels: A numpy array of ints or floats.
    neg_labels: A numpy array of ints or floats.
    split_value: An int.
                 If split_value is 300, the training set will consist of the first
                 300 emails in pos_emails plus the first 300 emails in neg_emails,
                 and the rest of the emails go into the test set.
                 
    Returns
    -------
    A 4-tuple of (X_train, X_test, y_train, y_test)
    All numpy arrays.
    '''
    
    # YOUR CODE HERE
    
    return X_train, X_test, y_train, y_test
```

```python
X_train, X_test, y_train, y_test = train_test_split(
    pos_emails, neg_emails, pos_labels, neg_labels, split_value=400
    )
```

<!-- grade_id: train_test_split_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(X_train, np.ndarray)
assert_is_instance(X_test, np.ndarray)
assert_is_instance(y_train, np.ndarray)
assert_is_instance(y_test, np.ndarray)
assert_array_equal(X_train[:400], pos_emails[:400])
assert_array_equal(X_train[400:], neg_emails[:400])
assert_array_equal(X_test[:len(pos_emails) - 400], pos_emails[400:])
assert_array_equal(X_test[len(pos_emails) - 400:], neg_emails[400:])
assert_array_equal(y_train[:400], pos_labels[:400])
assert_array_equal(y_train[400:], neg_labels[:400])
assert_array_equal(y_test[:len(pos_labels) - 400], pos_labels[400:])
assert_array_equal(y_test[len(pos_labels) - 400:], neg_labels[400:])
```

The following cell deletes the original two numpy email arrays to reduce memory requirements. 

```python
# Remove the original two numpy arrays
%xdel pos_emails
%xdel neg_emails
```

## Spam classification

- Use unigrams and bigrams,
- Build a pipeline by using TfidfVectorizer and LinearSVC,
- Name the first step tf and the second step svc,
- Use default parameters for both TfidfVectorizer and LinearSVC, and
- Use English stop words.

<!-- grade_id: fit_and_predict_answer, points: 0, grade: false, solution: true -->
```python
def fit_and_predict(X_train, y_train, X_test, random_state):
    '''
    Creates a document term matrix and uses SVM classifier to make document classifications.
    Uses English stop words.
    
    Parameters
    ----------
    X_train: A list of strings.
    y_train: A list of strings.
    X_test: A list of strings.
    random_state: A np.random.RandomState instance.
    
    Returns
    -------
    A tuple of (clf, y_pred)
    clf: A Pipeline instance.
    y_pred: A numpy array.    '''
    
    # YOUR CODE HERE
    
    return clf, y_pred
```

```python
clf, y_pred = fit_and_predict(X_train, y_train, X_test, random_state=check_random_state(0))
score = accuracy_score(y_test, y_pred)
print("SVC prediction accuracy = {0:5.1f}%".format(100.0 * score))
```

<!-- grade_id: fit_and_predict_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(clf, Pipeline)
assert_is_instance(y_pred, np.ndarray)
tf = clf.named_steps['tf']
assert_is_instance(tf, TfidfVectorizer)
assert_is_instance(clf.named_steps['svc'], LinearSVC)
assert_equal(tf.ngram_range, (1, 2))
assert_equal(tf.stop_words, 'english')
assert_equal(len(y_pred), len(y_test))
assert_array_equal(y_pred[:10], [1] * 10)
assert_array_equal(y_pred[-10:], [0] * 10)
assert_almost_equal(score, 1.0)
```