This notebook will be collected automatically at **6pm on Monday** from `/home/data_scientist/assignments/Week8` directory on the course JupyterHub server. If you work on this assignment on the course Jupyterhub server, just make sure that you save your work and instructors will pull your notebooks automatically after the deadline. If you work on this assignment locally, the only way to submit assignments is via Jupyterhub, and you have to place the notebook file in the correct directory with the correct file name before the deadline.

1. Make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).
2. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.
3. Do not change the file path or the file name of this notebook.
4. Make sure that you save your work (in the menubar, select `File` → `Save and CheckPoint`)

## Problem 8.3. Social Media: Web

- Find the titles of all decisions handed down by the U.S. Supreme Court in 2014.

The relelvant URL is [http://www.supremecourt.gov/opinions/slipopinion/14](http://www.supremecourt.gov/opinions/slipopinion/14).

```python
from IPython.display import HTML, display
import requests
import bs4
from bs4 import BeautifulSoup
import re
from lxml import html

from nose.tools import assert_equal, assert_is_instance, assert_true
```

```python
url = 'http://www.supremecourt.gov/opinions/slipopinion/14'
```

<!-- grade_id: find_titles_answer, points: 0, grade: false, solution: true -->
```python
def find_titles(url):
    '''
    Parses the US Supreme Court term opinions webpage
    and returns the titles of all decisions.
    
    Parameters
    ----------
    url: A string. Example: 'http://www.supremecourt.gov/opinions/slipopinion/14'
    
    Returns
    -------
    A list of strings
    '''
    
    # YOUR CODE HERE
    
    return titles
```

```python
titles_2014 = find_titles(url)
for title in titles_2014:
    print(title)
```

<!-- grade_id: find_titles_test, points: 20, grade: true, solution: false -->
```python
answer = [
    'Glossip v. Gross',
    'Arizona State Legislature v. Arizona Independent Redistricting Comm’n',
    'Michigan v. EPA',
    'Obergefell v. Hodges',
    'Johnson v. United States',
    'Texas Dept. of Housing and Community Affairs v. Inclusive Communities Project, Inc.',
    'King v. Burwell',
    'Kimble v. Marvel Entertainment, LLC',
    'Los Angeles v. Patel',
    'Kingsley v. Hendrickson',
    'Horne v. Department of Agriculture',
    'Brumfield v. Cain',
    'Davis v. Ayala',
    'Ohio v. Clark',
    'Walker v. Texas Div., Sons of Confederate Veterans, Inc.',
    'McFadden v. United States',
    'Reed v. Town of Gilbert',
    'Reyes Mata v. Lynch',
    'Baker Botts L.L.P. v. ASARCO LLC',
    'Kerry v. Din',
    'Zivotofsky v. Kerry',
    'Taylor v. Barkes',
    'Mellouli v. Lynch',
    'Bank of America, N. A. v. Caulkett',
    'EEOC v. Abercrombie & Fitch Stores, Inc.',
    'Elonis v. United States',
    'Wellness Int’l Network, Ltd. v. Sharif',
    'Kellogg Brown & Root Services, Inc. v. United States ex rel. Carter',
    'Commil USA, LLC v. Cisco Systems, Inc.',
    'Henderson v. United States',
    'City and County of San Francisco v. Sheehan',
    'Comptroller of Treasury of Md. v. Wynne',
    'Coleman v. Tollefson',
    'Tibble v. Edison Int’l',
    'Harris v. Viegelahn',
    'Bullard v. Blue Hills Bank',
    'Mach Mining, LLC v. EEOC',
    'Williams-Yulee v. Florida Bar',
    'United States v. Kwai Fun Wong',
    'Oneok, Inc. v. Learjet, Inc.',
    'Rodriguez v. United States',
    'Armstrong v. Exceptional Child Center, Inc.',
    'Woods v. Donald',
    'Grady v. North Carolina',
    'Alabama Legislative Black Caucus v. Alabama',
    'Young v. United Parcel Service, Inc.',
    'Omnicare, Inc. v. Laborers Dist. Council Constr. Industry Pension Fund',
    'B&B Hardware, Inc. v. Hargis Industries, Inc.',
    'Kansas v. Nebraska',
    'Perez v. Mortgage Bankers Assn.',
    'Department of Transportation v. Association of American Railroads',
    'Alabama Dept. of Revenue v. CSX Transp., Inc.',
    'Direct Marketing Assn. v. Brohl',
    'Yates v. United States',
    'North Carolina Bd. of Dental Examiners v. FTC',
    'Kansas v. Nebraska',
    'M&G Polymers USA, LLC v. Tackett',
    'Hana Financial, Inc. v. Hana Bank',
    'Gelboim v. Bank of America Corp.',
    'Department of Homeland Security v. MacLean',
    'Christeson v. Roper',
    'Holt v. Hobbs',
    'Teva Pharmaceuticals USA, Inc. v. Sandoz, Inc.',
    'T-Mobile South, LLC v. City of Roswell',
    'Jennings v. Stephens',
    'Whitfield v. United States',
    'Jesinoski v. Countrywide Home Loans, Inc.',
    'United States v. California',
    'Dart Cherokee Basin Operating Co. v. Owens',
    'Heien v. North Carolina',
    'Warger v. Shauers',
    'Integrity Staffing Solutions, Inc. v. Busk ',
    'Glebe v. Frost',
    'Carroll v. Carman',
    'Johnson v. City of Shelby',
    'Lopez v. Smith'
    ]
assert_is_instance(titles_2014, list)
assert_equal(len(titles_2014), 76)
assert_equal(set(titles_2014), set(answer))
```