This notebook will be collected automatically at **6pm on Monday** from `/home/data_scientist/assignments/Week12` directory on the course JupyterHub server. If you work on this assignment on the course Jupyterhub server, just make sure that you save your work and instructors will pull your notebooks automatically after the deadline. If you work on this assignment locally, the only way to submit assignments is via Jupyterhub, and you have to place the notebook file in the correct directory with the correct file name before the deadline.

1. Make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).
2. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.
3. Do not change the file path or the file name of this notebook.
4. Make sure that you save your work (in the menubar, select `File` → `Save and CheckPoint`)

# Problem 12.2. Apache Pig.

In this problem, we will run Pig to compute the average rating for each movie in the MovieLens data set.

```python
from nose.tools import assert_equal, assert_almost_equal
```

Since each data file in the MovieLens includes a header row, we employ the following Bash script to 
remove the first line from each file and make copies of the data files we will be using.

```python
%%bash

sed '1d' $HOME/data/ml-latest-small/ratings.csv > ratings.csv
sed '1d' $HOME/data/ml-latest-small/movies.csv > movies.csv

echo
echo '***** Ratings File *****'
head ratings.csv

echo
echo '***** Movies File *****'
head movies.csv
```

## Pig: Average

Write a Pig script that

- Imports `ratings.csv` and `movies.csv`,
- Groups all reviews by movie ID and uses [AVG](https://pig.apache.org/docs/r0.7.0/piglatin_ref2.html#AVG) to compute the average rating for each movie,
- Joins the two data sets on the movie ID column, and
- Uses the DUMP command to display the first 10 rows.

The resulting schema should contain five columns:

```
(movie ID from ratings.csv, average rating, movie ID from movies.csv, title, genre)
```

For example, the first line should be

```
(1,3.9073275862068964,1,Toy Story (1995),Adventure|Animation|Children|Comedy|Fantasy)
```

<!-- grade_id: average_answer, points: 0, grade: false, solution: true -->
```python
%%writefile average.pig

--YOUR CODE HERE
```

We run Pig in local mode and redirect the standard error output to a file.

```python
average_ratings = !pig -x local -f average.pig 2> pig_stderr.log
print('\n'.join(average_ratings))
```

To debug, uncomment and run the following code cell.

```python
#!cat pig_stderr.log
```

## Tests

<!-- grade_id: average_test, points: 30, grade: true, solution: false -->
```python
answer = [
    '(1,3.9073275862068964,1,Toy Story (1995),Adventure|Animation|Children|Comedy|Fantasy)',
    '(2,3.3532608695652173,2,Jumanji (1995),Adventure|Children|Fantasy)',
    '(3,3.189655172413793,3,Grumpier Old Men (1995),Comedy|Romance)',
    '(4,2.8181818181818183,4,Waiting to Exhale (1995),Comedy|Drama|Romance)',
    '(5,3.25,5,Father of the Bride Part II (1995),Comedy)',
    '(6,4.073913043478261,6,Heat (1995),Action|Crime|Thriller)',
    '(7,3.381818181818182,7,Sabrina (1995),Comedy|Romance)',
    '(8,3.6666666666666665,8,Tom and Huck (1995),Adventure|Children)',
    '(9,2.869565217391304,9,Sudden Death (1995),Action)',
    '(10,3.6,10,GoldenEye (1995),Action|Adventure|Thriller)'
    ]

a1 = [a.split(',') for a in answer]
a2 = [a.split(',') for a in average_ratings]

for irow, row in enumerate(answer):
    for icol in [0, 2, 3, 4]:
        assert_equal(a1[irow][icol], a2[irow][icol])
    # floating numbers in 1st column
    # I don't think it's necessary, but just in case
    assert_almost_equal(float(a1[irow][1]), float(a2[irow][1]))
```

## Cleanup

```python
%%bash
# Remove pig log files
rm -f pig*.log

# Remove our pig scripts
rm -f *.pig

# Remove csv files
rm movies.csv ratings.csv
```